# Exercise 3 - Levenshtein distance
(3 points)

Finish the `calcLevenshteinDistance` method by implementing the calculation of the Levenshtein distance of the two given Strings. Use dynamic programming as described in the lecture slides.

#### Notes

- Do not add additional external libraries.
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
public int calcLevenshteinDistance(String string1, String string2) {
    int distance = 0;
    // YOUR CODE HERE
    
    String s = string1, t = string2;
    
    if (s == null || t == null) {
          throw new IllegalArgumentException("Strings must not be null");
      }
    
    int n = s.length(); // length of s
      int m = t.length(); // length of t

      if (n == 0) {
          return m;
      } else if (m == 0) {
          return n;
      }

      if (n > m) {
          // swap the input strings to consume less memory
          String tmp = s;
          s = t;
          t = tmp;
          n = m;
          m = t.length();
      }

      int p[] = new int[n+1]; //'previous' cost array, horizontally
      int d[] = new int[n+1]; // cost array, horizontally
      int _d[]; //placeholder to assist in swapping p and d

      // indexes into strings s and t
      int i; // iterates through s
      int j; // iterates through t

      char t_j; // jth character of t

      int cost; // cost

      for (i = 0; i<=n; i++) {
          p[i] = i;
      }

      for (j = 1; j<=m; j++) {
          t_j = t.charAt(j-1);
          d[0] = j;

          for (i=1; i<=n; i++) {
              cost = s.charAt(i-1)==t_j ? 0 : 1;
              // minimum of cell to the left+1, to the top+1, diagonally left and up +cost
              d[i] = Math.min(Math.min(d[i-1]+1, p[i]+1),  p[i-1]+cost);
          }

          // copy current distance counts to 'previous row' distance counts
          _d = p;
          p = d;
          d = _d;
      }

      // our last action in the above loop was to switch d and p, so p now 
      // actually has the most recent cost counts
      distance = p[n];
    
    return distance;
}

// This line should make sure that compile errors are directly identified when executing this cell
// (the line itself does not produce any meaningful result)
Arrays.sort(new int[calcLevenshteinDistance("a","a")]);
```

# Evaluation

- Run the following cell to test your implementation.
- You can ignore the cells afterwards.

<!-- grade_id: cell-autograde-visible, points: 0, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

public void checkLevenshteinDistance(String s1, String s2, int expectedDistance) {
    try {
        int result = calcLevenshteinDistance(s1, s2);
        Assertions.assertEquals(expectedDistance, result, "For levenshtein(\"" + s1 + "\",\"" + s2
                + "\") your solution returned " + result + " while " + expectedDistance + " has been expected.");
        System.out.println("Test successfully completed.");
    } catch (AssertionFailedError e) {
        throw e;
    } catch (Throwable e) {
        System.err.println("Your solution caused an unexpected error:");
        throw e;
    }
}

checkLevenshteinDistance("abababab", "aaaaaaaa", 4);
checkLevenshteinDistance("rupert", "Robert", 3);
checkLevenshteinDistance("Robert", "rupert", 3);
checkLevenshteinDistance("Robert", "Robert", 0);
checkLevenshteinDistance("abcde", "", 5);
checkLevenshteinDistance("", "abcde", 5);


String word1 = "bqsvtzfhktnqovnjmkktluayihcfelhfqfsordvkxrmweenobehobnokrzabxcglfavxwyqlhjjlfxyx"
    + "ddptklfuwvtvwgecpdskvovhrrlaoqzjqkpauhpkkrxtrwmrevimykuyvstfgsmeegqjvxcstegiiqgl"
    + "auprgfuchymgvznavfpxnflxlujvtsshwhsnkunetfiwvqppaanugcovhnvwwedxywtkswmfeccatpna"
    + "fzntnagkjpsinosekeahzmqtcpjkaxzmunclbcurppwumfnmxlrfgjxsddblhiorwiaqcjlbnbfkdjfr"
    + "fqeqzrlfoiejfcmwdrcxneoquogbbuwdsykxvjpmfahmoybrqmneyrlutiunrbhfclzzhdtlxnwivoeg"
    + "zfzxmargbkxkjnqiuqsgakhzgxkxtchewatusacyeodaankidwezlopjabyzaxpxzffxnqaredklqron"
    + "slyzqgtfdjkaiwdcgalwvgpnlbxvuyqmhxsudcthfgziahugmjdhpqjucyvffviupnzmdtbatqbfewhy"
    + "jyfsijxxuyoyonldvclzrdoujbaskietubkfoopbeucqrgxzvmkwcdmbuglfcnkizkefskvmqazbtyyn"
    + "lrwptebsqvrtwtfukoesbjtgcxesvimzzdyoqvnclbedqhhgirhdyexourkwhzpjxjctdvqpiprfedsr"
    + "xywdzncqihboggtekrknjbrlflrfegxjbvaqqjcodbzmvgefbateguibcajueuohwftimarzgjjygyng"
    + "ayihcjnlnhidcooojummtbjuaddjsfyxqqpmphsuicxeeaizqrsrdomxhxwymjtfuigiijwowdoolbdv"
    + "bvwtipnybxndkubuzskeycpjuuiutokzlibbprywcnovjrbrhmlsoyumyxuaswbkoopaxbndrjyocqle"
    + "ishdwuuummzmtovjidjetgsfftrfmiqjckfphsvt";
String word2 = "bqsvtzfhktnuosnjmkktluayihcfelhfqfsordvkxrmweenobelobnocrzabxcglnavxcyqlhjjlfxyx"
    + "odptklfurvtvwiecpdrkvovhrrlnoizjqkpauhpkkrxtrwmrmvimykuyvstfgsmeenqjvxcstegiiqgl"
    + "aufrgfuchymgvznavfpxnflxlujvtssxwhsnkuneqfiwvqpdaanugcnvhavwwedxywtkswmfeccatpna"
    + "fzntnagkjpsinosekeahzmqtcpnkapzmunclbcurppwumfnmxmrfglxjddblhiorwiaqcjlbnbfkdjfw"
    + "fqeqzrjfoiejkcmwdrcxneonuogbmtwdsfkxvjpkfahmrynrqgneyrfutiunrbhfclzzhdtlxpwinocg"
    + "zfzxmargbkzkjnqiuqsgakfzgxkctbwewatusacyeodaankidwyzlopqabyzaxpxzffxnqaredkjqron"
    + "slyzqgtfdokaiwdcgalwvtpnlbxvuyqmhxsudcthfgzidhugmjdhpqjucyvffviupnzbdnbatqbfewhy"
    + "jyflcjxxuyoyonldvcpardoujbaskietubkfoopbeucqrgxzvmkwcdmbuglfqnkizkefhkvmqazbtkvn"
    + "lrwptebsqvrtwtfukoesbjtgcxesvimzzdpoqvnclbedqhomcrhdyexouukwwzpfxjctdiqpiprfedsr"
    + "xyodzncqijbogutekrknbbrlflrfegxjbvaqqjcodbzmvgefbateguiscajueuohjftimvrzgjjygymg"
    + "zyihcjnlnzidcaoojulmtbjuaddjslyxqlpmxhsuibxeeaizqrkrdomxhxwymjtfuigiijwowdoolbdv"
    + "svwtihnybxndkubhzskeycpjuuiutokzlibbprywcnovjrbrvllsogumyxuwswqkoouaxbnidjywcqle"
    + "ishbwuuummzmtovjidjetgsfftrfmiqjckfphsvt";
checkLevenshteinDistance(word1, word2, 99);
checkLevenshteinDistance(word2, word1, 99);
```

Output:
```
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
Test successfully completed.
```

<!-- grade_id: cell-autograde-1, points: 2, grade: true, solution: false -->
```java
// Ignore this cell
```

<!-- grade_id: cell-autograde-2, points: 1, grade: true, solution: false -->
```java
// Ignore this cell
```