# Viterbi Algorithm
(10 points)

Finish the implementation of the `ViterbiAlgorithm` class using the dynamic programming approach presented in the lecture. The class has a constructor taking the necessary information of a Hidden Markov Model and a `getStateSequence` method returning the most probable state of sequences for a given sequence of observations.

* The constructor takes the following parameters:
  * `String[] states` - an array containing the set of states $Q$ in the same order as in the transition matrix and the emission matrix.
  * `String[] observationVocab` - an array containing all possible observations $V$ in the order they have in the emission matrix.
  * `double[][] transitionMatrix` - the transition matrix $A$ where `transitionMatrix[i][j]` contains the probability $a_{ij}$ for a transition from $q_i$ to $q_j$ (since the `states` array starts at 0 and not at 1, `transitionMatrix[i][j]` gives the transition from `states[i - 1]` to `states[j - 1]`). Note that the transition matrix has two more states as described in the lecture slides. The start state always has the index `0` marks the start state in the matrix while the final state always has the index `transitionMatrix.length - 1`. 
  * `double[][] emissionMatrix` - the emission matrix $B$ containing the probabilities $b_{ik}$ that state $q_i$ emits the observation $v_k$.

* The `getStateSequence` method takes the following parameters:
  * `String[] observations` - an array containing the observations for which the most probable sequence of states should be returned.

* The method should return an instance of the `StateSequence` class. This class contains two following attributes:
  * `String[] states` - the sequence of states which emitted a certain sequence of observations.
  * `double logProbability` - the logarithm of the probability of the state sequence (including the emission probabilities for the sequence of observations).

#### Example

The visible tests rely on the ice cream example available in the lecture slides (see lecture slide 14 / page 22 in the PDF). The array of states and observations are given as follows:
```java
String[] states = new String[] { "HOT", "COLD" };
String[] observationVocab = new String[] { "1", "2", "3" };
```
double[][] transitionMatrix = new double[][] { { 0, 0.8, 0.2, 0 }, { 0, 0.6, 0.3, 0.1 }, { 0, 0.4, 0.5, 0.1 },
                { 0, 0, 0, 0 } };

The transition matrix $A$ taken from the automaton picture looks like the following table. (**Please note** that although the states `HOT` and `COLD` have the ids `0` and `1` in the `states` array, the have the ids `1` and `2` in the transition matrix.)

<table>
    <tr>
        <th>from \ to</th>
        <th><p align="center">start</p></th>
        <th><p align="center">HOT</p></th>
        <th><p align="center">COLD</p></th>
        <th><p align="center">end</p></th>
    </tr>
    <tr>
        <td><p align="center"><b>start</b></p></td>
        <td><p align="center">0</p></td>
        <td><p align="center">0.8</p></td>
        <td><p align="center">0.2</p></td>
        <td><p align="center">0</p></td>
    </tr>
    <tr>
        <td><p align="center"><b>HOT</b></p></td>
        <td><p align="center">0</p></td>
        <td><p align="center">0.6</p></td>
        <td><p align="center">0.3</p></td>
        <td><p align="center">0.1</p></td>
    </tr>
    <tr>
        <td><p align="center"><b>COLD</b></p></td>
        <td><p align="center">0</p></td>
        <td><p align="center">0.4</p></td>
        <td><p align="center">0.5</p></td>
        <td><p align="center">0.1</p></td>
    </tr>
    <tr>
        <td><p align="center"><b>end</b></p></td>
        <td><p align="center">0</p></td>
        <td><p align="center">0</p></td>
        <td><p align="center">0</p></td>
        <td><p align="center">0</p></td>
    </tr>
</table>

The emission matrix $B$ taken from the automaton picture looks like the following table. (Please note that this matrix does neither contain the start nor the final state since both are not emitting any observation.)

<table>
    <tr>
        <th>state \ observation</th>
        <th><p align="center">1</p></th>
        <th><p align="center">2</p></th>
        <th><p align="center">3</p></th>
    </tr>
    <tr>
        <td><p align="center"><b>HOT</b></p></td>
        <td><p align="center">0.2</p></td>
        <td><p align="center">0.4</p></td>
        <td><p align="center">0.4</p></td>
    </tr>
    <tr>
        <td><p align="center"><b>COLD</b></p></td>
        <td><p align="center">0.5</p></td>
        <td><p align="center">0.4</p></td>
        <td><p align="center">0.1</p></td>
    </tr>
</table>

#### Hints

- The input matrices will contain probabilities (not their logarithms)
- The test is separated into two different cells. The first cell will test some example observation sequences and compare it with an expected result. The second cell will do the same but for a larger, generated sequence.
- The number of states and observations are not limited to the ice cream example. For the hidden tests, we will use a different scenario with different observations and states.

#### Notes

- Do not add additional external libraries.
- As a final test run before submission
  - restart the kernel (wait for the kernel to be ready)
  - go through all cells from top to bottom and execute them
  - make sure that the whole process
    1. does not throw any errors and
    2. takes less than 2:00 minutes (the evaluation has a max runtime of 5 minutes per file including the hidden tests)
- Interface
  - You can use _[TAB]_ for autocompletion and _[SHIFT]_+_[TAB]_ for code inspection.
  - Use _Menu_ -> _View_ -> _Toggle Line Numbers_ for debugging.
  - Check _Menu_ -> _Help_ -> _Keyboard Shortcuts_.
- Finish
  - Save your solution by clicking on the _disk icon_.
  - Finally, choose _Menu_ -> _File_ -> _Close and Halt_.
  - Do not forget to _Submit_ your solution in the _Assignments_ view.

<!-- grade_id: cell-code, points: 0, grade: false, solution: true -->
```java
/**
 * Simple structure for storing a sequence of states and its probability.
 */
public static class StateSequence {
    /**
     * The sequence of states.
     */
    public final String[] states;
    /**
     * The logarithm of the probability of this sequence.
     */
    public final double logProbability;

    public StateSequence(String[] states, double logProbability) {
        this.states = states;
        this.logProbability = logProbability;
    }
}

/**
 * A class implementing the Viterbi algorithm based on an Hidden Markov Model
 * with the given states, observations, transition matrix and emission matrix.
 */
public static class ViterbiAlgorithm {
    // YOUR CODE HERE

	String[] states;
	String[] sequence;
	double[][] transitionMatrix;
	String[] observationVocab;
	double[][] emissionMatrix;

	/**
	 * Constructor.
	 */
	public ViterbiAlgorithm(String[] states, String[] observationVocab, double[][] transitionMatrix,
			double[][] emissionMatrix) {
		// YOUR CODE HERE

		this.states = states;
		this.emissionMatrix = emissionMatrix;
		this.observationVocab = observationVocab;
		this.transitionMatrix = transitionMatrix;

	}

	/**
	 * Returns the sequence of states which has the highest probability to create
	 * the given sequence of observations.
	 * 
	 * @param observations a sequence of observations
	 * @return the sequence of states
	 */
	public StateSequence getStateSequence(String[] observations) {
		StateSequence sequence = null;
		// YOUR CODE HERE
		int T = observations.length;
		int N = states.length;

		double[][] vitrebi = new double[N + 2][T];
		Integer[][] backPointer = new Integer[N + 2][T];
		
		int id = 0;
		for(int k =0; k<observationVocab.length; k++)
			if(observations[0] == observationVocab[k]) id = k;

		for (int s = 1; s <= N; s++) {

			vitrebi[s][0] = Math.log(transitionMatrix[0][s] * emissionMatrix[s-1][id]);
			backPointer[s][0] = 0;
		}

		List<String> trail = new ArrayList<String>();
        
		// populate Viterbi table with max. probabilities
		for (int t = 1; t < T; t++) {

			int index = 1;
			for(int k =0; k<observationVocab.length; k++)
				if(observations[t] == observationVocab[k]) index = k;

			for (int s = 1; s <= N; s++) {
				
				double probMax = -Double.MAX_VALUE;
				double bpMax = -Double.MAX_VALUE;

				int argMax = 0;
				//System.out.println("t and s values are  - "+t+ " " +s);
				for (int i = 1; i <=N; i++) {
                    double calcProb = vitrebi[i][t - 1] + Math.log(transitionMatrix[i][s]);
					if (calcProb >= probMax) {
						probMax = calcProb;
						argMax = i;   //newly added
					}

				}

				vitrebi[s][t] = Math.log(emissionMatrix[s-1][index]) + probMax;
				backPointer[s][t] = argMax;

			}
		}
		
		double probMax = -Double.MAX_VALUE;
		int argMax = 0;
		for (int s = 1; s <= N; s++) {
            
			double calcProb = vitrebi[s][T-1] + Math.log(transitionMatrix[s][N + 1]);
			if (calcProb >= probMax) {
				probMax = calcProb;
				argMax = s;
			}

		}
		vitrebi[N+1][T-1] = probMax;
		backPointer[N + 1][T-1] = argMax;


		// return the backtrace path by following states from backPointer[qF,T] backwards
		List<String> path = new ArrayList<String>();

		
		
		for (int i = 1; i < T; i++) {
			double max = vitrebi[1][i];
			int sid=1;
			for(int s=1; s<=N; s++)
			{
				if(vitrebi[s][i] >= max)
				{
					max = vitrebi[s][i];
					sid = backPointer[s][i];
				}

			}

			path.add(states[sid-1]);

		}
		
		path.add(states[backPointer[N+1][T-1] - 1]);
		

		String[] stringArray = path.toArray(new String[0]);
		sequence = new StateSequence(stringArray, vitrebi[N + 1][T - 1]);
		
		return sequence;
	}
}

// This line should make sure that compile errors are directly identified when executing this cell
// (the line itself does not produce any meaningful result)
new ViterbiAlgorithm(new String[0],new String[0],new double[2][2],new double[0][0]);
System.out.println("compiled");
```

Output:
```
compiled
```

# Evaluation

- Run the following cell to test your implementation.
- You can ignore the cells afterwards.

<!-- grade_id: cell-autograde-visible, points: 0, grade: true, solution: false -->
```java
%maven org.junit.jupiter:junit-jupiter-api:5.3.1
import org.junit.jupiter.api.Assertions;
import org.opentest4j.AssertionFailedError;

public static final double DELTA = 0.000001;

public static void checkViterbi(String[] states, double[][] transitionMatrix, String[] observationVocab,
        double[][] emissionMatrix, String[] observations, StateSequence expectedSequence) {
    try {
        ViterbiAlgorithm viterbi = new ViterbiAlgorithm(states, observationVocab, transitionMatrix, emissionMatrix);
        long time1 = System.currentTimeMillis();
        StateSequence sequence = viterbi.getStateSequence(observations);
        time1 = System.currentTimeMillis() - time1;
        System.out.println("Viterbi took " + time1 + "ms");
        Assertions.assertArrayEquals(expectedSequence.states, sequence.states);
        double diff = Math.abs(expectedSequence.logProbability - sequence.logProbability);
        Assertions.assertTrue(diff < DELTA, "The calculated probability (" + sequence.logProbability
                + ") does not match the expected probability (" + expectedSequence.logProbability + ").");
        System.out.println("Test passed");
    } catch (AssertionFailedError e) {
        throw e;
    } catch (Throwable e) {
        System.err.println("Your solution caused an unexpected error:");
        throw e;
    }
}

String observations[];
StateSequence expectedSequence;
String[] states;
String[] sequence;
double[][] transitionMatrix;
String[] observationVocab;
double[][] emissionMatrix;

System.out.println("---------- Ice cream example ----------");
states = new String[] { "HOT", "COLD" };
transitionMatrix = new double[][] { { 0, 0.8, 0.2, 0 }, { 0, 0.6, 0.3, 0.1 }, { 0, 0.4, 0.5, 0.1 },
        { 0, 0, 0, 0 } };
observationVocab = new String[] { "1", "2", "3" };
emissionMatrix = new double[][] { { 0.2, 0.4, 0.4 }, { 0.5, 0.4, 0.1 } };
observations = new String[] { "3", "1", "3" };
expectedSequence = new StateSequence(new String[] { "HOT", "HOT", "HOT" }, Math.log(0.0009216));
checkViterbi(states, transitionMatrix, observationVocab, emissionMatrix, observations, expectedSequence);

observations = new String[] { "3", "2", "1", "1" };
expectedSequence = new StateSequence(new String[] { "HOT", "HOT", "COLD", "COLD" }, Math.log(0.000288));
checkViterbi(states, transitionMatrix, observationVocab, emissionMatrix, observations, expectedSequence);

observations = new String[] { "1", "3", "3", "2", "3", "2", "1", "3", "1", "1", "1" };
expectedSequence = new StateSequence(
        new String[] { "HOT", "HOT", "HOT", "HOT", "HOT", "HOT", "HOT", "HOT", "COLD", "COLD", "COLD" },
        Math.log(3.439853568E-9));
checkViterbi(states, transitionMatrix, observationVocab, emissionMatrix, observations, expectedSequence);

```

Output:
```
---------- Ice cream example ----------
Viterbi took 0ms
Test passed
Viterbi took 0ms
Test passed
Viterbi took 0ms
Test passed
```

<!-- grade_id: cell-autograde-visible-2, points: 0, grade: true, solution: false -->
```java
System.out.println("---------- Ice cream example ----------");
observations = new String[1000];
Arrays.fill(observations, "3");
sequence = new String[1000];
Arrays.fill(sequence, "HOT");
expectedSequence = new StateSequence(sequence,
        Math.log(0.8) + (999 * Math.log(0.6)) + (1000 * Math.log(0.4) + Math.log(0.1)));
checkViterbi(states, transitionMatrix, observationVocab, emissionMatrix, observations, expectedSequence);
```

Output:
```
---------- Ice cream example ----------
Viterbi took 2ms
Test passed
```

<!-- grade_id: cell-autograde-1, points: 6, grade: true, solution: false -->
```java
// Ignore this cell
```

<!-- grade_id: cell-autograde-2, points: 4, grade: true, solution: false -->
```java
// Ignore this cell
```